# This is the main input for unit testing.
# Only this file is passed as command line argument.

rule_files:
    - 12_prometheusrules.yaml

evaluation_interval: 1m

tests:
    # Test 1.
    - interval: 1m
      # Series data.
      input_series:
          - series: 'vsphere_esxi_version_total{api_version="6.5",version=""}'
            values: '1 1 2 1 1 2 2 2 2 2 2 2 2 2 2'
          - series: 'vsphere_esxi_version_total{api_version="7.0.3.0",version="7.0.3"}'
            values: '3 3 3 3 3 3 3 3 3 3 3 3 3 3 3'


      # Unit test for alerting rules.
      alert_rule_test:
          # Unit test 1.
          - eval_time: 10m
            alertname: VSphereOlderVersionPresent
            exp_alerts:
                # Alert 1.
                - exp_labels:
                      severity: info
                      api_version: "6.5"
                  exp_annotations:
                      summary: "Detected vSphere host with ESXi version less than 7.0.2 in Openshift cluster."
                      description: |
                        The cluster is using ESXi hosts which are on version less than 7.0.2, which is being deprecated by Openshift. A future version of
                        Openshift will remove support for ESXi version less than 7.0.2 and it is recommended to update your hosts to the latest ESXi version.
